# Review module

**Instructions**

In order to complete this review module, we recommend you follow these instructions:

1. Complete the code cells provided to you in this notebook, but do **not** change the names of the variables. If you do that, the autograder will fail and you will not receive any points.
2. Run all the answered cells before you run the testing cells. The answers must exist before they are graded!
3. Remove from each cell the code `raise NotImplementedError()` and replace it with your implementation.
4. Do not round any quantity.

### Exercise 1 (1 point)

The `wine` DataFrame contains data for about 1,600 Portuguese red wine samples, including quality ratings assigned by human tasters (the `quality` column) and chemical markers obtained via lab tests:

```python
import pandas as pd
import statsmodels.formula.api as sm
wine = pd.read_csv("data/winequality-red.csv", sep=";")
wine.head()
```

Output:
```
fixed acidity  volatile acidity  citric acid  residual sugar  chlorides  \
0            7.4              0.70         0.00             1.9      0.076   
1            7.8              0.88         0.00             2.6      0.098   
2            7.8              0.76         0.04             2.3      0.092   
3           11.2              0.28         0.56             1.9      0.075   
4            7.4              0.70         0.00             1.9      0.076   

   free sulfur dioxide  total sulfur dioxide  density    pH  sulphates  \
0                 11.0                  34.0   0.9978  3.51       0.56   
1                 25.0                  67.0   0.9968  3.20       0.68   
2                 15.0                  54.0   0.9970  3.26       0.65   
3                 17.0                  60.0   0.9980  3.16       0.58   
4                 11.0                  34.0   0.9978  3.51       0.56   

   alcohol  quality  
0      9.4        5  
1      9.8        5  
2      9.8        5  
3
```

Using `statsmodels`, code this simple linear regression:

$$
quality = \beta_0 + \beta_1 alcohol + \varepsilon
$$

Save the **fitted** model in a variable called `reg1`.

<!-- grade_id: cell-12cba7ae22a27e20, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

formula1 = 'quality ~ alcohol'
model1 = sm.ols(formula = formula1, data = wine)
reg1 = model1.fit()
print(reg1.summary())
#raise NotImplementedError() # Remove this line when you enter your solution
```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:                quality   R-squared:                       0.227
Model:                            OLS   Adj. R-squared:                  0.226
Method:                 Least Squares   F-statistic:                     468.3
Date:                Wed, 01 Jun 2022   Prob (F-statistic):           2.83e-91
Time:                        06:58:54   Log-Likelihood:                -1721.1
No. Observations:                1599   AIC:                             3446.
Df Residuals:                    1597   BIC:                             3457.
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

### Exercise 2 (1 point)

Using `statsmodels`, code this multiple linear regression:

$$
quality = \beta_0 + \beta_1 alcohol + \beta_2 pH + \beta_3 sulphates + \beta_4 density + \beta_5 residual{\ }sugar + \beta_6 citric{\ }acid + \varepsilon
$$

Save the **fitted** model in a variable called `reg2`.

**Hint:** `statsmodels` formulas expect variable names that don't have spaces in them. To include a variable name with spaces, use [`Q` syntax](https://patsy.readthedocs.io/en/latest/builtins-reference.html#patsy.builtins.Q) (`Q` stands for "quote") like this: `Q("total sulfur dioxide")`

<!-- grade_id: cell-d1e86a6b82bf94b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
formula2 = 'quality ~ alcohol+pH+sulphates+density+Q("residual sugar")+Q("citric acid")'
model2 = sm.ols(formula = formula2, data = wine)
reg2 = model2.fit()
print(reg2.summary())
```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:                quality   R-squared:                       0.289
Model:                            OLS   Adj. R-squared:                  0.286
Method:                 Least Squares   F-statistic:                     107.6
Date:                Wed, 01 Jun 2022   Prob (F-statistic):          5.38e-114
Time:                        07:00:46   Log-Likelihood:                -1654.4
No. Observations:                1599   AIC:                             3323.
Df Residuals:                    1592   BIC:                             3361.
Df Model:                           6                                         
Covariance Type:            nonrobust                                         
=======================================================================================
                          coef    std err          t      P>|t|      [0
```

### Exercise 3 (1 point)

Complete the code in the below cell to extract the following values from the model you fitted in the previous exercise:

* The $R^2$ (save it in a variable called `rsq`)
* The intercept $\beta_0$ (save it in a variable called `beta0`)
* The coefficient of the $alcohol$ variable, $\beta_1$ (save it in a variable called `beta1`)

**Note:** Do not change the name of the function.

<!-- grade_id: cell-89c7e70a8af98a29, points: 0, grade: false, solution: true -->
```python
def extract_results(fitted_model):
    """
    Extracts R squared, intercept and beta 1 from the model.
    """
    # YOUR CODE HERE
    rsq=reg2.rsquared
    beta0=reg2.params.values[0]
    beta1=reg2.params.values[1]

    return rsq, beta0, beta1

#rsq, beta0, beta1 = extract_results(reg2)
#print("R squared: ", rsq)
#print("Intercept: ", beta0)
#print("Beta 1: ", beta1)
```

## Testing cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get an error when checking your answers.

<!-- grade_id: cell-644427ae97585a8f, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert "reg1" in globals(), "Ex. 1 - You need to create the reg1 variable!"
assert type(reg1).__name__ == 'RegressionResultsWrapper', "Ex. 1 - The reg1 variable is not a statsmodels regression model! Did you forget to fit it? Remember to use the .fit() method!"
print("Exercise 1 passed the preliminary sanity check.")
```

Output:
```
Exercise 1 passed the preliminary sanity check.
```

<!-- grade_id: cell-967d7c9f19f2866e, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert "reg2" in globals(), "Ex. 2 - You need to create the reg2 variable!"
assert type(reg2).__name__ == 'RegressionResultsWrapper', "Ex. 2 - The reg2 variable is not a statsmodels regression model! Did you forget to fit it? Remember to use the .fit() method!"
assert len(reg2.params) == 7, "Ex. 2 - Your regression has either too many or too few input variables! There must be seven betas: One intercept and six input variables"
print("Exercise 2 passed the preliminary sanity check.")
```

Output:
```
Exercise 2 passed the preliminary sanity check.
```

<!-- grade_id: cell-42cea078db9cfba8, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert "extract_results" in globals(), "Ex. 3 - You need to create the extract_results function! Did you forget to run the cell where you defined it?"
print("Exercise 3 passed the preliminary sanity check.")
```

Output:
```
Exercise 3 passed the preliminary sanity check.
```

## Attribution

"Wine quality", October 07, 2009, Cortez, Paulo *et al.*, Creative Commons Attribution 4.0 International license, https://archive-beta.ics.uci.edu/ml/datasets/wine+quality